# Problem Set 02

## POLI 175 - Machine Learning for Social Scientists

### Due Date: 02/22/2023

In this problem set we will play with the Civil Conflict dataset. You can find it in [here](https://github.com/umbertomig/POLI175public/tree/main/data/mshk-pa-2017).

The file you are going to use for this PS is the `SambanisImp.csv`. The code book is also on the folder.

The full paper compares Random Forests with Logistic Regressions to predict civil conflict. You can find it in here: [full paper](https://doi.org/10.1093/pan/mpv024)

# 1. Loading Packages

Load:

- pandas
- numpy
- matplotlib.pyplot
- seaborn
- sklearn (I'd suggest pretty much all the packs we used in class)

Alias them in the standard way. Warning: do not load the `DecisionBoundaryDisplay`.

<!-- grade_id: cell-ad2435cff0a8cfd9, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-743491bf46f29312, points: 0.5, grade: true, solution: false -->
```python
# Testing
assert 'np' in dir()
assert 'pd' in dir()
assert 'plt' in dir()
assert 'sns' in dir()

```

# 2. Loading Data

Load the dataset and save it in an object called `dat`.

<!-- grade_id: cell-c8dc3385e93dd279, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-afd5d022b58f6a54, points: 1, grade: true, solution: false -->
```python
assert type(dat) == pd.core.frame.DataFrame

```

# 3. Subsetting

The are too many variables, some of them not useful. We are going to use the same as the used in [Hegre and Sambanis (2006)](https://journals.sagepub.com/doi/pdf/10.1177/0022002706289303). Subset your data to only keep these variables.

```python
## Sambanis vars
sambvars = ["warstds", "ager", "agexp", "anoc", "army85", "autch98", 
            "auto4","autonomy", "avgnabo", "centpol3", "coldwar", 
            "decade1", "decade2","decade3", "decade4", "dem", "dem4", 
            "demch98", "dlang", "drel", "durable", "ef", "ef2", "ehet", 
            "elfo", "elfo2", "etdo4590", "expgdp", "exrec", "fedpol3", 
            "fuelexp", "gdpgrowth", "geo1", "geo2", "geo34", "geo57", 
            "geo69", "geo8", "illiteracy", "incumb", "infant", "inst", 
            "inst3", "life", "lmtnest", "ln_gdpen", "lpopns", "major", 
            "manuexp", "milper", "mirps0", "mirps1", "mirps2", "mirps3", 
            "nat_war", "ncontig", "nmgdp", "nmdp4_alt", "numlang", 
            "nwstate", "oil", "p4mchg", "parcomp", "parreg", "part", 
            "partfree", "plural", "plurrel", "pol4", "pol4m", "pol4sq", 
            "polch98", "polcomp", "popdense", "presi", "pri", "proxregc", 
            "ptime", "reg", "regd4_alt", "relfrac", "seceduc", "second", 
            "semipol3", "sip2", "sxpnew", "sxpsq", "tnatwar", "trade",
            "warhist", "xconst"]

## Target variable
target = "warstds"

## Predictors
predictors = ["ager", "agexp", "anoc", "army85", "autch98", 
            "auto4","autonomy", "avgnabo", "centpol3", "coldwar", 
            "decade1", "decade2","decade3", "decade4", "dem", "dem4", 
            "demch98", "dlang", "drel", "durable", "ef", "ef2", "ehet", 
            "elfo", "elfo2", "etdo4590", "expgdp", "exrec", "fedpol3", 
            "fuelexp", "gdpgrowth", "geo1", "geo2", "geo34", "geo57", 
            "geo69", "geo8", "illiteracy", "incumb", "infant", "inst", 
            "inst3", "life", "lmtnest", "ln_gdpen", "lpopns", "major", 
            "manuexp", "milper", "mirps0", "mirps1", "mirps2", "mirps3", 
            "nat_war", "ncontig", "nmgdp", "nmdp4_alt", "numlang", 
            "nwstate", "oil", "p4mchg", "parcomp", "parreg", "part", 
            "partfree", "plural", "plurrel", "pol4", "pol4m", "pol4sq", 
            "polch98", "polcomp", "popdense", "presi", "pri", "proxregc", 
            "ptime", "reg", "regd4_alt", "relfrac", "seceduc", "second", 
            "semipol3", "sip2", "sxpnew", "sxpsq", "tnatwar", "trade",
            "warhist", "xconst"]
```

<!-- grade_id: cell-efc94562934e65e6, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-9f72db4416a18450, points: 1, grade: true, solution: false -->
```python
assert dat.columns[1] == 'ager'

```

# 4. Linear Regression

Your target variable in here is `warstds`, which means whether a country is experiencing conflict or not. I saved it in an object called `target`. All the predictors are called `predictors`. Run a Linear Regression using [scikit learn](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html#sklearn.linear_model.LinearRegression) to try to explain conflict based on all the predictors. Then, find the predicted outcomes and recode them with the following heuristics:

- 0 if prediction below or equal 0.5
- 1 if prediction above 0.5

Make a confusion matrix of your results. Do you have a good prediction model? Explain.

**Important**: you should rewrite the `linreg` with the fitting model, and the `y_pred_linreg` with the predictions.

```python
## Creating the target variable and the predictors' matrix
y = dat[target]
X = dat[predictors]
```

<!-- grade_id: cell-ef386fd8186f0431, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-11a7211754415b1f, points: 1, grade: true, solution: false -->
```python
assert type(linreg) == LinearRegression
assert y_pred_linreg[0] == 0

```

<!-- grade_id: cell-614973634f868f3b, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# 5. Logistic Regression

Fit the same model using a logistic regression. Use the parameter `solver = 'liblinear'`. Make a confusion matrix of your results. Do you have a good prediction model? Explain.

**Important**: you should rewrite the `logreg` with the fitting model, and the `y_pred_logreg` with the predictions.

<!-- grade_id: cell-725c289ed0019e12, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-3bc391ae62c7f1db, points: 1, grade: true, solution: false -->
```python
assert type(logreg) == LogisticRegression
assert y_pred_logreg[0] == 0

```

<!-- grade_id: cell-370da1ad5fc3bc51, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# 6. Linear Discriminant Analysis

Fit the same model using a linear discriminant analysis. Make a confusion matrix of your results. Do you have a good prediction model? Explain.

**Important**: you should rewrite the `lda` with the fitting model, and the `y_pred_lda` with the predictions.

<!-- grade_id: cell-899c718edc78d014, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-50ab2a652d5ff773, points: 1, grade: true, solution: false -->
```python
assert type(lda) == LinearDiscriminantAnalysis
assert y_pred_lda[140] != y[140]

```

<!-- grade_id: cell-f7dbdaf67381d463, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# 7. Quadratic Discriminant Analysis

Fit the same model using a linear discriminant analysis. Make a confusion matrix of your results. Do you have a good prediction model? Explain. (Special attention to the warning message. What does that mean? Check online!)

**Important**: you should rewrite the `qda` with the fitting model, and the `y_pred_qda` with the predictions.

<!-- grade_id: cell-98e429995cd94630, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-ddbf9b4e13e32f29, points: 1, grade: true, solution: false -->
```python
assert type(qda) == QuadraticDiscriminantAnalysis
assert y_pred_qda[0] == 0

```

<!-- grade_id: cell-499e24ee2af8b004, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# 8. Naïve Bayes Classifier

Fit the same model using the Gaussian Naïve Bayes classifier. Make a confusion matrix of your results. Do you have a good prediction model? Explain.

**Important**: you should rewrite the `nbayes` with the fitting model, and the `y_pred_naiveb` with the predictions.

<!-- grade_id: cell-4e66290f198526b8, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-68d285c8286c85cd, points: 1, grade: true, solution: false -->
```python
assert type(nbayes) == GaussianNB
assert y_pred_naiveb[0] == 0

```

<!-- grade_id: cell-9448f00053041451, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE